Maximum and minimum

<!-- BEGIN QUESTION -->

Please complete the function below that return the element with minimum value in a given list. You need to complete the functions using three different loop techniques: `while`, `for`, and `for + range` respectively. 

Note that you will **NOT** get marks if you use the built-in Python function such as `sort()`, `min()`, or `max()`.

```python
import math

def min_for(li): # for loop
    # INSERT YOUR CODE BELOW
    # ~ 4 lines
    if len(li) == 0:
        return None
    
    m = math.inf
    for num in li:
        if num < m:
            m = num 
    return m


def min_while(li): # while loop
    # INSERT YOUR CODE BELOW
    # ~ 4 lines
    if len(li) == 0:
        return None
    
    m = math.inf
    i = 0
    while i < len(li):
        if li[i] < m:
            m = li[i]
        i += 1
    return m

def min_range(li): # for + range
    # INSERT YOUR CODE BELOW
    # ~ 4 lines
    if len(li) == 0:
        return None

    m = math.inf
    for i in range(len(li)):
        if li[i] < m:
            m = li[i]
    return m


# ================
# TEST CASES BELOW
def test(li):
    v1 = min_for(li)
    v2 = min_while(li)
    v3 = min_range(li)
    print(f'{v1} : {v2} : {v3}')

li = [3, 4, 2, 15, 17, 30, 5, 45]
test(li) # 2 : 2 : 2

li = [5, 8, 12, -5, -7]
test(li) #-7 : -7 : -7
```

Output:
```
2 : 2 : 2
-7 : -7 : -7
```

palindrome check

A **palindrome** is a word, phrase, number or sequence of words that reads the same backward as forward. Please follow the instructions below, in sequence, to complete your palindrome check function.

### Task 2.1 

Complete the function `reverse_string` function that takes one string and return its reverse.

```python
def reverse_string(s):
    ns=''    
    for c in s:
        # COMPLETE THE CODE BELOW
        # ~ 1 LINE        
        ns = c + ns
    return ns


#################
#while loop
#def reverse_string(s):
#    ns = ''
#    index = len(s) - 1
    index = 0
#    
#    while index >= 0:
#        ns += s[index]
#        index -= 1

            index += 1
#    return ns


# ================
# TEST CASES BELOW
s = reverse_string('test.')
print(s) # '.tset'
```

Output:
```
.tset
```

### Task 2.2

Please write a `compare_string` function that compare if two strings are identical. The spec of the function is
- it has two parameters `str1` and `str2`. You can assume `str1` and `str2` have same length.
- it returns `-1` if the content of str1 is identical to str2 (as a practice, please compare both string character by character.)
- it returns an integer `n`, where `n` is the index of the first character that is different between `str1` and `str2`

<font color='red'>Note, you should not use built-in string comparison methods. Instead, your code must iterate through each characters and compare</font>

```python
def compare_string(str1, str2):
    r = -1
    # INPUT YOUR CODE BELOW
    # Iterate through each characters and compare
    # ~ 6 LINES
    
    # for i in range(len(str1)):
    #     c1 = str1[i]
    #     c2 = str2[i]
    #     if c1 != c2:
    #         r = i
    #         break
    # # END SOLUTION
    # return r
        
 
##############################
def compare_string(str1, str2):
    r = -1
    i = 0
    
    while i < len(str1) and i < len(str2):
        if str1[i] != str2[i]:
            r = i
            break
        i += 1
    
    # if r == -1:
    #     if len(str1) != len(str2):
    #         r = len(str1) if len(str1) < len(str2) else len(str2)
    
    return r
        
# ================
# TEST CASES BELOW
a = compare_string('test', 'test')
print(a) # -1
a = compare_string('xest', 'test') # first char is different
print(a) # 0
a = compare_string('text', 'test') # 3rd char is different
print(a) #                         
```

Output:
```
-1
0
2
```

### Task 2.3 

Please write a **is_palindrome** function that
- takes one string parameter
- return True if the string is a palindrome and return False otherwise

```python
#a word that is the same when you read it forwards from the beginning or backwards from the end
#e.g. level, kayak
def is_palindrome(s):
    # INSERT YOUR CODE BELOW
    # USE THE ABOVE TWO FUNCTIONS ABOVE TO COMPLETE THE FUNCTION
    # ~ 6 LINES
    
    
    
    # step 1 - s = anna | alice 
    s2 = reverse_string(s)
    
    #step 2 - compare s and s2 anna | ecilac 
    rs = reverse_string(s)
    r = compare_string(s, rs) s, s1 anna and anna
    
    if r == -1:
        return True
    else:
        return False
    
p = is_palindrome('anna') # True
print(p)
p = is_palindrome('civic') # True
print(p)
p = is_palindrome('alice') # False
print(p)
```

Output:
```
True
True
False
```

Caesar cipher 

The Caesar cipher is one of the earliest known and simplest ciphers. It is a type of substitution cipher in which each letter in the plaintext is 'shifted' a certain number of places down the alphabet. For example, with a shift of 2, the word "hi", will become "jk". Because **h** is replaced by **j**, **i** is replaced by **k**.

In this task, you will write a function that encrypts a message with Caesar Cipher and decrypt it back into plaintext. For simplicity, we can assume that the user will only use lowercase letters. 

Below are some helper code for you. the 
- char_list containing all 26 lowercase character, 
- the `.index` function returns the index of a value in the list.

```python
# HELPER CODE
char_list = ['a', 'b', 'c', 'd', 'e', 'f', 'g', 'h', 'i', 'j', 'k', 'l', 'm', 'n', 'o', 'p', 'q', 'r', 's', 't', 'u', 'v', 'w', 'x', 'y', 'z']

# below is a quick example of shifting the character by 2
idx = char_list.index('h')          #get the index of character 'h'
idx2 = (idx + 2 ) % len(char_list)  #shift the index by 2
new_char = char_list[idx2]          #get the encrypted character using the new index

print(f'h -> {new_char}')
```

Output:
```
h -> j
```

### Task 3.1 Encrypt

Once you understand the code above, please complete the `caesar_encrypt` function below, the function has two parameters
- the input string. Note that the input string might contain space character
- `n`, the number of shift 

```python
# def caesar_encrypt(s, n):
#     enc = '' # the encrypted string

#     # COMPLETE THE CODE BELOW
#     # iterate through each character and shift it by n, you should skip the space character
#     # ~ 7 lines
#     for c in s:
#         if c != ' ':
#             idx = char_list.index(c)          
#             idx2 = (idx + n ) % len(char_list)  
#             enc = enc + char_list[idx2] 
#         else:
#             enc = enc + c
#     return enc

#very classic question in C ROT13
#https://en.wikipedia.org/wiki/ROT13
#https://www.w3resource.com/python-exercises/string/python-data-type-string-exercise-25.php
###############
def caesar_encrypt(s, n):
    enc = '' # the encrypted string
    i = 0
    while i < len(s):
        c = s[i]
        if c != ' ':
            idx = char_list.index(c)
            idx2 = (idx + n) % len(char_list) #26
            enc = enc + char_list[idx2]
        else:
            enc = enc + c // enc += c
        i += 1
    
    return enc

if item == ' ':
    enc = enc = item
else 
# TEST CASES BELOW
es = caesar_encrypt('an apple a day', 4)
print(es) # er ettpi e hec

es = caesar_encrypt('keeps the doctor away', 8)
print(es) # smmxa bpm lwkbwz ieig
```

Output:
```
er ettpi e hec
smmxa bpm lwkbwz ieig
```

### Task 3.2 Decrypt
Following Task 1.1, when the receiver received the encrypted message, a decryptor is needed to translate the encrypted word into a plain one. Please write a program that decrypt the encrypted message.


```python
def caesar_decrypt(s, n):
    # COMPLETE THIS FUNCTION
    # HINT: Use the caesar_encrypt but with a different shift :)
    # ~1 LINE 
    return caesar_encrypt(s, -n)

# TEST CASE BELOW
n = 12
s = 'an apple a day keeps the doctor away'
es = caesar_encrypt(s, n)
print(es) # mz mbbxq m pmk wqqbe ftq paofad mimk
ds = caesar_decrypt(es, n)
print(ds) # an apple a day keeps the doctor away
```

Output:
```
mz mbbxq m pmk wqqbe ftq paofad mimk
an apple a day keeps the doctor away
```

---

Tests for Task 3.1 (points: 15)
```
>>> caesar_encrypt('this is a testing string', 5) == 'ymnx nx f yjxynsl xywnsl'
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> caesar_encrypt('adelaide', 23) == "xabixfab"
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Task1 (points: 20)


Tests for Task2.1 (points: 10)
```
>>> reverse_string('an apple') == 'elppa na'
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> reverse_string('doctor away') == 'yawa rotcod'
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Task2.2 (points: 10)
```
>>> compare_string('happy', 'happy') == -1
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> compare_string('habby', 'happy') == 2
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> compare_string('hapBy', 'happy') == 3
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Task2.3 (points: 10)
```
>>> is_palindrome('apple') == False
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> is_palindrome('civic') == True
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> is_palindrome('kayak') == True
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Task3.2 (points: 15)
```
>>> caesar_decrypt('xabixfab', 23) == 'adelaide'
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> caesar_decrypt('ixxtm', 8) == 'apple'
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
